# Lab 08: Circular, Array-backed Queue

## Overview

For this assignment you will implement a circular, array-backed queue data structure.

In the following class, which you are to complete, the backing array will be created and populated with `None`s in the `__init__` method, and the `head` and `tail` indexes set to sentinel values (you shouldn't need to modify `__init__`). Enqueueing and Dequeueing items will take place at the tail and head, with `tail` and `head` tracking the position of the most recently enqueued item and that of the next item to dequeue, respectively. To simplify testing, your implementation should make sure that when dequeuing an item its slot in the array is reset to `None`, and when the queue is emptied its `head` and `tail` attributes should be set to `-1`.

Because of the fixed size backing array, the `enqueue` operation is defined to raise a `RuntimeError` when the queue is full — the same exception should be raised when `dequeue` is called on an empty queue.

Finally, the `resize` method will allow the array underlying the queue to be increased in size. It is up to you how to implement this (you can either leave the elements in their current positions, though this may require "unwrapping" elements, or you can simply move all elements towards the front of the array). You may assume that `resize` will only be called with a value greater than the current length of the underlying array.

<!-- grade_id: queue, points: 0, grade: false, solution: true -->
```python
class Queue:
    def __init__(self, limit=10):
        self.data = [None] * limit
        self.head = -1
        self.tail = -1

    #def __next_slot__(self, int): #call this with._next_slot
    #    return (int + 1) % len(self)
        
    def enqueue(self, val):
        print(self.data)
        print(self.head)
        print(self.tail)
        if self.head == -1:
            self.data[0] = val
            self.head = self.tail = 0
        else:
            temp = ((self.tail +1) % len(self.data))
            if self.data[temp] is None:
                self.tail = temp
                self.data[self.tail] = val
            elif temp == self.head:
                if self.data[temp] is None:
                    self.tail = self.head = 0
                else:
                    raise RuntimeError
        print(self.data)   
            
            #if temp == self.head:
            #    raise RuntimeError
            #else:
            #    self.tail = temp
            #    self.data[self.tail] = val
            
                
            #if self.data[self.tail] is None:
                #raise RuntimeError
            #    self.data[self.tail] = val
            #elif self.data[self.tail] == self.data[self.head]:
            #if self.tail == self.head:
            #    raise RuntimeError
                
            #self.data[self.tail] = val
            
            #if self.data.count(None) == 0:
            #    raise RuntimeError
            #if self.data.count(None) == 0:
            #    raise RuntimeError
            #if self.tail & self.head == None:
            #    raise RuntimeError
        
    def dequeue(self):
        if self.head == -1:
            raise RuntimeError
        else:
            val = self.data[self.head]
            self.data[self.head] = None
            self.head = (self.head + 1) % len(self.data)
            if self.data[self.head] == None:
                self.head = -1
                #yield self.empty
            return val
            #if self.head == self.tail:
                
    
    def resize(self, newsize):
        assert(len(self.data) < newsize)
        
        print(self.data)
        print('boo')
        
        #newSlots = [None] * newsize
        #self.data[self.tail].append(newSlots)
        a = self.head
        b = self.tail
        #a = len(self.data)
        for i in range(0, newsize):
            self.data.insert(b + 1, None)
            self.head = (self.head + 1) % len(self.data)
            print(self.data)
        #self.head = a
        #self.tail = b
        print(self.head)
        print(b)
            
        print(self.data)
        #oldList = self.data
        #newList = [None] * newsize
        #self.data = oldList + newList
        #newList.head = -1
        #newList.tail = -1
        #for i in (oldList):
        #    enqueue(self.data, oldList[i])
            
        #for i in range(newsize):
        #    self.data.append
        #newSlots = newsize - len(self.data) #5
        #self.data.append([None] * newSlots)
        
        
        
    
    def empty(self):
        return self.head == -1
    
    def __bool__(self):
        return not self.empty()
    
    def __str__(self):
        if not(self):
            return ''
        return ', '.join(str(x) for x in self)
    
    def __repr__(self):
        return str(self)
    
    def __iter__(self):
        #for i in self.data:
        #    return self.data[i]
        for i in range(len(self.data)):
            if self.data[self.head] is not None:
                yield self.data[self.head]
            self.head = (self.head + 1) % len(self.data)
        
           
        
```

<!-- grade_id: queue_tests_1, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()

q = Queue(5)
tc.assertEqual(q.data, [None] * 5)

for i in range(5):
    q.enqueue(i)
    
with tc.assertRaises(RuntimeError):
    q.enqueue(5)

for i in range(5):
    tc.assertEqual(q.dequeue(), i)
    
tc.assertTrue(q.empty())
```

Output:
```
[None, None, None, None, None]
-1
-1
[0, None, None, None, None]
[0, None, None, None, None]
0
0
[0, 1, None, None, None]
[0, 1, None, None, None]
0
1
[0, 1, 2, None, None]
[0, 1, 2, None, None]
0
2
[0, 1, 2, 3, None]
[0, 1, 2, 3, None]
0
3
[0, 1, 2, 3, 4]
[0, 1, 2, 3, 4]
0
4
```

<!-- grade_id: queue_tests_2, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()

q = Queue(10)

for i in range(6):
    q.enqueue(i)
    
tc.assertEqual(q.data.count(None), 4)

for i in range(5):
    q.dequeue()
    
tc.assertFalse(q.empty())
tc.assertEqual(q.data.count(None), 9)
tc.assertEqual(q.head, q.tail)
tc.assertEqual(q.head, 5)

for i in range(9):
    q.enqueue(i)

with tc.assertRaises(RuntimeError):
    q.enqueue(10)

for x, y in zip(q, [5] + list(range(9))):
    tc.assertEqual(x, y)
    
tc.assertEqual(q.dequeue(), 5)
for i in range(9):
    tc.assertEqual(q.dequeue(), i)

tc.assertTrue(q.empty())
```

Output:
```
[None, None, None, None, None, None, None, None, None, None]
-1
-1
[0, None, None, None, None, None, None, None, None, None]
[0, None, None, None, None, None, None, None, None, None]
0
0
[0, 1, None, None, None, None, None, None, None, None]
[0, 1, None, None, None, None, None, None, None, None]
0
1
[0, 1, 2, None, None, None, None, None, None, None]
[0, 1, 2, None, None, None, None, None, None, None]
0
2
[0, 1, 2, 3, None, None, None, None, None, None]
[0, 1, 2, 3, None, None, None, None, None, None]
0
3
[0, 1, 2, 3, 4, None, None, None, None, None]
[0, 1, 2, 3, 4, None, None, None, None, None]
0
4
[0, 1, 2, 3, 4, 5, None, None, None, None]
[None, None, None, None, None, 5, None, None, None, None]
5
5
[None, None, None, None, None, 5, 0, None, None, None]
[None, None, None, None, None, 5, 0, None, None, None]
5
6
[None, None, None, None, None, 5, 0, 1, None, None]
[None, None, None, None, None, 5, 0, 1, None, None]
5
7
[None, None, None, None, None, 5, 0, 1, 2, None]
[None, None, None
```

<!-- grade_id: queue_tests_3, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()

q = Queue(5)
for i in range(5):
    q.enqueue(i)
for i in range(4):
    q.dequeue()
for i in range(5, 9):
    q.enqueue(i)
    
with tc.assertRaises(RuntimeError):
    q.enqueue(10)

q.resize(10)

for x, y in zip(q, range(4, 9)):
    tc.assertEqual(x, y)
    
for i in range(9, 14):
    q.enqueue(i)

for i in range(4, 14):
    tc.assertEqual(q.dequeue(), i)
    
tc.assertTrue(q.empty())
tc.assertEqual(q.head, -1)
```

Output:
```
[None, None, None, None, None]
-1
-1
[0, None, None, None, None]
[0, None, None, None, None]
0
0
[0, 1, None, None, None]
[0, 1, None, None, None]
0
1
[0, 1, 2, None, None]
[0, 1, 2, None, None]
0
2
[0, 1, 2, 3, None]
[0, 1, 2, 3, None]
0
3
[0, 1, 2, 3, 4]
[None, None, None, None, 4]
4
4
[5, None, None, None, 4]
[5, None, None, None, 4]
4
0
[5, 6, None, None, 4]
[5, 6, None, None, 4]
4
1
[5, 6, 7, None, 4]
[5, 6, 7, None, 4]
4
2
[5, 6, 7, 8, 4]
[5, 6, 7, 8, 4]
4
3
[5, 6, 7, 8, 4]
boo
[5, 6, 7, 8, None, 4]
[5, 6, 7, 8, None, None, 4]
[5, 6, 7, 8, None, None, None, 4]
[5, 6, 7, 8, None, None, None, None, 4]
[5, 6, 7, 8, None, None, None, None, None, 4]
[5, 6, 7, 8, None, None, None, None, None, None, 4]
[5, 6, 7, 8, None, None, None, None, None, None, None, 4]
[5, 6, 7, 8, None, None, None, None, None, None, None, None, 4]
[5, 6, 7, 8, None, None, None, None, None, None, None, None, None, 4]
[5, 6, 7, 8, None, None, None, None, None, None, None, None, None, None, 4]
14
3
[5, 6, 7, 8, Non
```

```python

```

```python

```

```python

```