# Lab 06: Ordered Hashtable

## Overview

For this assignment you will update and complete the implementation of the hashtable data structure presented in class, which exposes an API mirroring that of the built-in Python `dict`. When iterating over its contents (supported by the `__iter__`, `keys`, `values`, and `items` methods), your updated implementation will also reflect the order in which key/value pairs were originally inserted into the hashtable. This will require that you implement the two-tiered list system described during lecture.

The operations you will implement are listed alongside their descriptions below (`h` refers to a hashtable):

| Operation | Description |
|-----------|-------------|
| `h[k]`&nbsp;`=`&nbsp;`v` | If `h` does not contain key `k`, a new `k`&rightarrow;`v` mapping is added, else the value for key `k` is updated to `v`. |
| `h[k]`    | If `h` contains key `k`, the corresponding value is returned, else a `KeyError` is raised. |
| `del`&nbsp;`h[k]` | If `h` contains key `k`, it is removed along with its value, else a `KeyError` is raised. Note that if `k` is re-inserted at some later point it is considered a new key (for ordering purposes). |
| `k`&nbsp;`in`&nbsp;`h` | Returns `True` if key `k` is in `h`. |
| `len(h)` | Returns the number of keys in `h`. |
| `iter(h)` | Returns an iterator over all the keys in `h`, in the order they were added. |
| `h.keys()` | (Same as above) |
| `h.values()` | Returns an iterator over all the values in `h`, in the order they were added. |
| `h.items()` | Returns an iterator over all the key/value pairs (as tuples) in `h`, in the order they were added. |

Your hashtable will be provided with the initial number of buckets on creation (i.e., in `__init__`); your implementation must heed this value, as there may be performance ramifications if it does not.

<!-- grade_id: hashtable, points: 0, grade: false, solution: true -->
```python
class OrderedHashtable:
    class Node:
        """This class is used to create nodes in the singly linked "chains" in
        each hashtable bucket."""
        def __init__(self, index, next=None):
            # don't rename the following attributes!
            self.index = index
            self.next = next
        
    def __init__(self, n_buckets=1000):
        # the following two variables should be used to implement the "two-tiered" 
        # ordered hashtable described in class -- don't rename them!
        self.indices = [None] * n_buckets
        self.entries = []
        self.count = 0
        
    def __getitem__(self, key):
        hashidx = hash(key) % len(self.indices)
        hashnode = self.indices[hashidx]
        if hashnode is None:
            raise KeyError
        while hashnode is not None: #originally hashnode.next
            #originally nde = nde.next is here
            if key == self.entries[hashnode.index][0]:
                return self.entries[hashnode.index][1]
            hashnode = hashnode.next
        raise KeyError
        #hashIndex = hash(key) % len(self.indices)
        #entry = self.indices[hashIndex]
        #if entry:
        #    return entry[1]
        #else:
        #    raise KeyError()
    
    def __setitem__(self, key, val):
        
        hashidx = hash(key) % len(self.indices)
        node = self.indices[hashidx]
        while node:
            if key == self.entries[node.index][0]:
                self.entries[node.index][1] = val
                return
            node = node.next
        new = OrderedHashtable.Node(len(self.entries), self.indices[hashidx])
        self.indices[hashidx] = new
        self.entries.append([key,val])
        self.count += 1
        
    
    def __delitem__(self, key):
        
        idx = hash(key) % len(self.indices)
        n = self.indices[idx]
        newindex = -1
        if n and self.entries[n.index][0] == key:
            newindex = n.index
            self.indices[idx] = self.indices[idx].next
        else:
            while n.next:
                if self.entries[n.next.index][0] == key:
                    newindex = n.next.index
                    n.next = n.next.next
                    break
                n = n.next
        if newindex == -1:
            raise KeyError
        self.entries[newindex] = None
        self.count -=1
        #if current.next is not None:
        #    idx = current.next
        ##nextNode = current.next
        #while current is True:
        #    if self.entries[current.index][0] == key:
        #        self.entries[current.index] = None
        #        #self.count -= 1
        #        
        #        self.indices[hashidx] = self.indices[hashidx].next
        #        
        #        
        #    #idx = current.next.index
        #    elif self.entries[idx.index][0] == key:
        #        self.entries[idx.index] = None
        #        #self.count -=1
        #        
        #        current = self.indices[hashidx].next.next
        #        
        #    if idx or idx.next == None:
        #        raise KeyError
        #        break
        #    current = current.next
        #    idx = idx.next
        #    self.count-=1
        #raise KeyError
        #prior = self.indices[hashidx]
        #if prior:
        #    current = prior.next
        #while prior:
        #    if key == self.entries[prior.index][0]:
        #        self.entries[prior.index] = None
        #        self.count -= 1
        #    #if prior == current:
        #    #    self.entries[prior.index] = None
        #        if current.next:
        #            prior.index = current.index
        #            prior = prior.next
        #            current = current.next
        #            
        #        elif current.next is None:
        #            prior.index = None
        #    elif key == self.entries[current.index][0]:
        #        continue
        
        #raise KeyError
        
        #for i in self.entries:
        #    if i[0] == key:
        #        del key 
        
    def __contains__(self, key):
        try:
            _ = self[key]
            return True
        except:
            return False
        
    def __len__(self):
        return self.count
    
    def __iter__(self):
        for i in self.entries:
            if i is None:
                continue
            yield i[0]
        
    def keys(self):
        return iter(self)
    
    def values(self):
        for i in self.entries:
            if i is None:
                continue
            yield i[1]
                
    def items(self):
        for i in self.entries:
            if i is None:
                continue
            yield (i[0],i[1])
        
    def __str__(self):
        return '{ ' + ', '.join(str(k) + ': ' + str(v) for k, v in self.items()) + ' }'
            
    def __repr__(self):
        return str(self)
    
                    #prior.index = current.index
                #prior.next = current.next
            #prior = prior.next
            #current = current.next
            #if current.next:
                #current = current.next
            
```

<!-- grade_id: tests_short, points: 3, grade: true, solution: false -->
```python
# (3 tests) Short tests

from unittest import TestCase
import random

tc = TestCase()

ht = OrderedHashtable(2)

for k, v in (('batman', 'bruce wayne'), ('superman', 'clark kent'), ('spiderman', 'peter parker')):
    ht[k] = v
    
tc.assertEqual(len(ht), 3)

tc.assertEqual(ht['superman'], 'clark kent')

tc.assertTrue('spiderman' in ht)
tc.assertFalse('iron man' in ht)

with tc.assertRaises(KeyError):
    ht['iron man']
```

<!-- grade_id: tests_basic, points: 3, grade: true, solution: false -->
```python
# (3 points) Basic tests (insertion, fetch, count, chain-lengths)

from unittest import TestCase
import random

tc = TestCase()

class MyInt(int):
    def __hash__(self):
        """MyInts hash to themselves — already current Python default, 
        but just to ensure consistency."""
        return self
    
def ll_len(l):
    """Returns the length of a linked list with head `l` (assuming no sentinel)"""
    c = 0
    while l:
        c += 1
        l = l.next
    return c
    
ht = OrderedHashtable(10)
for i in range(25):
    ht[MyInt(i)] = i*2

tc.assertEqual(len(ht), 25)

for i in range(5):
    tc.assertEqual(ll_len(ht.indices[i]), 3)
    
for i in range(5, 10):
    tc.assertEqual(ll_len(ht.indices[i]), 2)

for i in range(25):
    tc.assertTrue(MyInt(i) in ht)
    tc.assertEqual(ht[MyInt(i)], i*2)
```

<!-- grade_id: tests_update, points: 3, grade: true, solution: false -->
```python
# (3 points) Update testing

from unittest import TestCase
import random

tc = TestCase()

ht = OrderedHashtable(100)
d = {}

for i in range(100):
    k, v = str(i), str(i*2)
    d[k] = v
    ht[k] = v
    
for j in range(0, 100, 2):
    k, v = str(i), str(i*3)
    d[k] = v
    ht[k] = v
    
for j in range(0, 100, 4):
    k, v = str(i), str(i*4)
    d[k] = v
    ht[k] = v
    
for i in range(100):
    tc.assertTrue(k in ht)
    tc.assertEqual(d[k], ht[k])
```

<!-- grade_id: tests_delete, points: 3, grade: true, solution: false -->
```python
# (3 points) Deletion testing

from unittest import TestCase
import random

tc = TestCase()

ht = OrderedHashtable(100)
d = {}

for i in range(100):
    k, v = str(i), str(random.randrange(10000000, 99999999))
    d[k] = v
    ht[k] = v

for _ in range(50):
    k = str(random.randrange(100))
    if k in d:
        del d[k]
        del ht[k]

tc.assertEqual(len(ht), len(d))

for k,v in ht.items():
    tc.assertEqual(d[k], v)
```

<!-- grade_id: tests_ordering, points: 4, grade: true, solution: false -->
```python
# (4 points) Iteration order testing

from unittest import TestCase
import random

tc = TestCase()

ht = OrderedHashtable(1000)
l = [str(i) for i in range(0, 1000)]
random.shuffle(l)

for x in l:
    ht[x] = x

for _ in range(50):
    idx_to_del = random.randrange(len(l))
    val_to_del = l[idx_to_del]
    del ht[val_to_del]
    del l[idx_to_del]
    if random.randrange(2) == 0:
        l.append(val_to_del)
        ht[val_to_del] = val_to_del

for x, y in zip(l, ht):
    tc.assertEqual(x, y)
```

<!-- grade_id: tests_stress, points: 4, grade: true, solution: false -->
```python
# (4 points) Stress testing

from unittest import TestCase
from time import time
import random

tc = TestCase()

ht = OrderedHashtable(100000)
d = {}

start = time()

for _ in range(100000):
    k, v = str(random.randrange(100000)), str(random.randrange(10000000, 99999999))
    d[k] = v
    ht[k] = v
    
for k,v in d.items():
    tc.assertTrue(k in ht)
    tc.assertEqual(d[k], ht[k])
    
end = time()
print(end-start)
tc.assertLess(end-start, 1.5, 'Your implementation ran too slow!')
```

Output:
```
1.1554291248321533
```

```python

```

```python

```